<div class = "alert alert-warning" align="center"> 
    <img src="https://www.egd.mg/wp-content/uploads/2020/01/cropped-logo-home-110px.png" align="left" width="290"  margin="50px"> 
    <center><h1>Lycée français de Tananarive</h1> 
        <h2> Jupyter Hub,  Spécialité NSI </h2></center>
</div>

Parcours séquentiel de listes
======
Codage des algorithmes de parcours en Python
------



# Préambule
L'objectif de cette partie est de comprendre comment fonctionnent ces algorithmes de parcours, pour cela il conviendra de n'utiliser que des opérations de boucle (`while` ou `for`) et des opérations conditonnelles (`if`).

# Algorithmes de cumul.
## Calcul de moyenne.
> Ecrire une fonction `moyenne(T)` qui prend en argument une liste d'entiers, et en renvoie la moyenne.

<!-- grade_id: cell-6d627ee1f74fb3d3, points: 0, grade: false, solution: true -->
```python
def moyenne(T) :
    res = 0
    for x in T:
        res += x
        c = res / len(T)
    return c
moyenne([10,20,13,14,8,11])
```

Output:
```
12.666666666666666
```

<!-- grade_id: cell-6ff4d860517f1279, points: 1, grade: true, solution: false -->
```python
assert moyenne([1,2,3]) == 2.0
assert moyenne([10,20,13,14,8,11]) == 12.666666666666666
assert moyenne([10]) == 10.0
```

> En modifiant la fonction `moyenne(T)`, écrire une fonction `compte_moyenne(T)` qui compte le nombre d'affectations réalisées lors du calcul de la moyenne


<!-- grade_id: cell-0a46f5fba3e37125, points: 0, grade: false, solution: true -->
```python
def moyenne(T) :
    res = 0
    for x in T:
        res += x
        c = res / len(T)
    return c

def compte_moyenne(T) :
    x = moyenne(T)
    y = len(T)
    return x, y
compte_moyenne([x ** 2 for x in range(50)])
```

Output:
```
(808.5, 50)
```

<!-- grade_id: cell-da3b7c1d50cf39bb, points: 2, grade: true, solution: false -->
```python
assert compte_moyenne([10]) == (10.0, 1)
assert compte_moyenne([i for i in range(100)]) == (49.5, 100)
assert compte_moyenne([i**2 for i in range(50)]) == (808.5, 50)
```

## Cacul de la variance
En statistique, on est amené à calculer la **variance** d'une série. 
> Ecrire une fonction `var_moy(T)` qui renvoie la moyenne et la variance des valeurs d'une liste d'entiers passée en paramétre. Essayez de n'utiliser qu'une seule boucle.

Le calcul de la variance  s'effectue ainsi :
- On calcule la somme des carrées des valeurs
- On divise cette cette somme par le nombre d'éléments de la liste, on calcule ainsi la moyenne des carrés des valeurs
- On calcule la moyenne simple des valeurs et on l'éléve au carré.
- On obtient la variance en soustrayant à la moyenne des carrés le carré de la moyenne.

$ variance = \frac{1}{n}\sum \limits_{i=0}^n {T_i}^2 -\left( \frac{1}{n} \sum \limits_{i=0}^n {T_i} \right)^2$

<!-- grade_id: cell-ba3263da1bcd1c7c, points: 0, grade: false, solution: true -->
```python
def somme_sum(T):
    v = 0
    for x in T:
        v += x
    return v

def var_moy(T):
    v = moyenne(T)
    sum_c = somme_sum([x ** 2 for x in T])
    mcv = sum_c / len(T)
    ms = moyenne(T) ** 2
    res = mcv - ms
    return v, res
```

<!-- grade_id: cell-41a98b72a10dcb0f, points: 3, grade: true, solution: false -->
```python
assert var_moy([1,1,1]) == (1.0, 0.0)
assert var_moy([1,0,2]) == (1.0, 0.6666666666666667)
assert var_moy([0,0,3]) == (1.0, 2.0)
assert var_moy([i for i in range(50)]) == (24.5, 208.25)
```

## Comptage d'éléments.


> Ecrire une fonction `pairs(T)` qui compte le nombre d'entiers pairs contenus dans une liste.

<!-- grade_id: cell-1bb47c827f06c832, points: 0, grade: false, solution: true -->
```python
def pairs(T) :
    cpt = 0
    for x in T:
        if x == 0:
            continue
        elif x % 2 == 0:
            cpt += 1
    return cpt
```

<!-- grade_id: cell-c93e2898d80dba28, points: 1, grade: true, solution: false -->
```python
assert pairs([1,1,1]) == 0
assert pairs([2]) == 1
assert pairs([2,1,2,1,2]) == 3
assert pairs([2,1,0,1,2]) == 2
```